# Lab: Birthday

## Happy Birthday!

Just kidding, it's *(probably, with P=0.9973)* not your birthday, but it might be someone's birthday in this room! Taylor Swift says she will come and give a free concert on THE QUAD on December 13th if it is someone's birthday on that day. Now obviously the students want to make this happen, so they ask you to do some calculations regarding the probability of the concert happening.

- For these puzzles, we assume that each year is **365 days** (ignore leap years)
- We also assume that everyone has an **equal chance** of a birthday each day *(not actually true in real life)*

## Group
Edit the next Python cell to add information about who you're working with in your lab section:

```python
partner1_name = ""
partner1_netid = ""
partner1_major = ""

partner2_name = ""
partner2_netid = ""
partner2_major = ""

partner3_name = ""
partner3_netid = ""
partner3_major = ""
```

## Table of Contents <a id='toc'></a>

**Use the Outline view (View > Command Palette > Type "outline") for navigation.**

* Puzzle 1
    * Puzzle 1.1
    * Puzzle 1.2
    * Puzzle 1.3
    * Puzzle 1.4
* Puzzle 2
    * Puzzle 2.1
    * Puzzle 2.2
    * Puzzle 2.3 
* Puzzle 3
    * Puzzle 3.1
    * Puzzle 3.2
    * Puzzle 3.3
* Puzzle 4
    * Puzzle 4.1
    * Puzzle 4.2
    * Puzzle 4.3
    * Puzzle 4.4
    * Puzzle 4.5
* Finale

## Puzzle 1: My Birthday

The **birthday problem** is a classic statistical problem that asks, *"How likely is it for at least two people in a room of **n** people to share a birthday?"*  In other words, what's the chance that there's at least one match of birthdays in a room with **n** people? We'll cover this in Puzzle \#2.


#### "My Birthday Problem"

To begin, let's consider **your birthday**.  We are actually going to start with a related, but very different question...

*How likely is it for at least one person in a room of **`n`** people to have the same birthday as **me**?*

### Puzzle 1.1: Sharing My Birthday with 1, 2, or 3 Other People

#### One Other Person (`n=1`)

If you are in a room with **one other person**, what is the probability you both have the same birthday? (*Of course you know your own birthday!*) There is no need for any python code for this puzzle, it should be a mathematical equation! Replace the ... with your answer in the following puzzles.

Hint (double-click this cell to see it): <span style="color:#ffffff00">Check Puzzle 1.2 for the formula of these probabilities.</span>

<!-- grade_id: cell-16021a893dda10f8, points: 0, grade: false, solution: true -->
```python
p_one_person_my_birthday = ...
p_one_person_my_birthday
```

#### Two Other People (`n=2`)

If you are in a room with **two other people**, what is the probability that at least one of them has the same birthday as you?

<!-- grade_id: cell-d776448804b67dfc, points: 0, grade: false, solution: true -->
```python
p_two_others_my_birthday = ...
p_two_others_my_birthday
```

#### Three Other People (`n=3`)

If you are in a room with **three other people**, what is the probability that at least one of them has the same birthday as you?

<!-- grade_id: cell-9b6271534fe56ab1, points: 0, grade: false, solution: true -->
```python
p_three_others_my_birthday = ...
p_three_others_my_birthday
```

<!-- grade_id: cell-2213109ad4022435, points: 0, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 1a ==
# - This read-only cell contains test cases for your previous cells.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.
import math
assert(math.isclose(p_one_person_my_birthday, 0.0027397260273972603)), "The probability for one birthday doesn't seem right."
assert(math.isclose(p_two_others_my_birthday, 0.005471945956089352)), "The probability for two birthdays doesn't seem right."
assert(math.isclose(p_three_others_my_birthday, 0.008196680350730179)), "The probability for three birthdays doesn't seem right."

## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
```

### Puzzle 1.2: Sharing My Birthday with `n` Other People

We have included a Python library developed just for DISCOVERY called `birthday` that provides the ability to run code related to birthdays.

This library works just like `pandas` in that:

1. You must `import birthday` to use the functions found in `birthday`.
2. We will introduce various "birthday functions" for you to use!


#### Puzzle 1: Birthday Function

The first birthday function calculates the probability that someone in the room shares **your birthday**!  From Puzzle 1.1, you found that this follows the following equation when you have **n other people in the room**:

$$
P(n) = 1 - \left(\frac{364}{365}\right)^n.
$$

For one other person ($n=1$):

\begin{align*}
P(n) &= 1 - \left(\frac{364}{365}\right)^1 \\
& = \frac{1}{365} \\
& \approx 0.0027397.
\end{align*}

For two other person ($n=2$):

\begin{align*}
P(n) &= 1 - \left(\frac{364}{365}\right)^2 \\
& \approx 0.0054719.
\end{align*}

We can do this with the `birthday` function: `birthday.myBirthday(n)` where `n` is how many other people are in the room.

#### Your Turn!

1. Import `birthday`
2. Using the `birthday.myBirthday(n)` function for **2** other people in the room

<!-- grade_id: cell-9fcc75505fada2fa, points: 0, grade: false, solution: true -->
```python
...
```

### Running The Function:

Fix the following calls to `birthday.myBirthday(n)` to call the function with the correct value of `n`.  Make sure the answers are what you expect.

```python
print('A room with one other person:', birthday.myBirthday(...))
print('A room with two other people:', birthday.myBirthday(...))
print('A room with three other people:', birthday.myBirthday(...))
print('A room with 100 other people:', birthday.myBirthday(...))
print('A room with 1000 other people:', birthday.myBirthday(...))
print('A room with 10000 other people:', birthday.myBirthday(...))
```

### Puzzle 1.3: Birthday DataFrame

A second birthday function provided is `birthday.myBirthday_DataFrame(n)`.  This function returns a DataFrame containing `n` rows with the probability you share a birthday with someone when there are `n` other people in the room.
`
- This is **very similar** to `pd.read_csv(url)` in that the both return a DataFrame.  This time, Python is generating the DataFrame instead of loading it from a file!

Use `birthday.myBirthday_DataFrame(n)` to create a DataFrame of birthdays for up to **2000** people in the same room and output it below. Store the result in `df`.

```python
...
```

Using your DataFrame `df`, we can create a line chart to show how the probability grows over time. Run the code below!


```python
df.plot.line(x="Other People", y="P(n)", grid=True)
```

### Puzzle 1.4: ❓ Individual Reflection Question ❓ Chance of Sharing Your Birthday vs. Two People Sharing a Birthday

 Is it more or less likely that someone in your lab shares **your** birthday **OR** that there's at least one shared birthday between **anyone** in your lab?

***Double click this cell and type your answer here.***

## Puzzle 2: The Classic Birthday Problem

The **birthday problem** is a classic problem that asks: *How likely is it for at least two people in a room to share a birthday?*<br>
**NOTE**: *Since we're no longer using your birthday, all birthdays here are random and not fixed.*

### Puzzle 2.1: At Least Two People Sharing a Birthday at a 1, 2, 3, or 4-Person Party

#### A One-Person Party (`n=1`)

Let's start with just one person in a room.  The probability that they share a birthday with someone else is... well... 0%.

$P(\mathrm{shared\ birthday\ with\ 1\ person}) = 0$


#### A Two-Person Party (`n=2`)

The probability of two people in a room sharing the same birthday is non-zero!  This is the probability of the second person having exactly the same birthday as the first ... 1 out of 365!

$P(\mathrm{shared\ birthday\ with\ 2\ people}) = \frac{1}{365}$


#### A Three-Person Party (`n=3`)

Now it gets interesting!  In a three person party, what is the probability that there is **at least one** match of birthdays?

Use the following cell to find the answer in Python, the solutions will be mathematical equations:

**HINT:** (double-click this cell to see it): <span style="color:#ffffff00">Probability of at least 1 is the same as saying 1 - Probability of no matches!</span>

<!-- grade_id: cell-9db7fd67d49459bd, points: 0, grade: false, solution: true -->
```python
p_three_person_party = ...
p_three_person_party
```

<!-- grade_id: cell-3aa5b5a4afef9fbc, points: 0, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 2 (with n=3) ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.
import math
assert(math.isclose(p_three_person_party, 0.008204165884781345)), "The probability doesn't seem right."

## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
```

#### A Four-Person Party (`n=4`)

In a four person party, what is the probability that there is **at least one** match of birthdays?

Use the following cell to find the answer in Python:

<!-- grade_id: cell-1a5b32da97b97164, points: 0, grade: false, solution: true -->
```python
p_four_person_party = ...
p_four_person_party
```

<!-- grade_id: cell-571231fcdfa9de08, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 2 (with n=4) ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.

assert(math.isclose(p_four_person_party, 0.016355912466550215)), "The probability doesn't seem right."

## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
```

#### Birthday Library: At Least Two People Sharing a Birthday at an `n`-Person Party

We now have a **third** birthday function `birthday.sharedBirthday(n)` that calculates the probability that at a party with **n** people, **at least two people** share a birthday.

Check if `birthday.sharedBirthday(4)` matches with your answer above:

<!-- grade_id: cell-8826e853f191d14f, points: 0, grade: false, solution: true -->
```python
...
```

Fix the following calls to `birthday.sharedBirthday(n)` to call the function with the correct value of `n`.  Make sure the answers are what you expect.

```python
print('A room with just you:',  birthday.sharedBirthday(...))
print('A room with two people:',  birthday.sharedBirthday(...))
print('A room with three people:',  birthday.sharedBirthday(...))
print('A room with four people:',  birthday.sharedBirthday(...))
print('A room with 40 people:',  birthday.sharedBirthday(...))
print('A room with 80 people:',  birthday.sharedBirthday(...))
print('A room with 400 people:',  birthday.sharedBirthday(...))
```

### Puzzle 2.2: Result Plot

Just like in Puzzle 1, the `birthday` library has a function for you to create a DataFrame containing the probabilities up to `n` total people at your party: `birthday.sharedBirthday_DataFrame(n)`!

Use `birthday.sharedBirthday_DataFrame(n)` to create a DataFrame of birthdays for up to **400** people in the same room and output it below:

```python
df = ...
df
```

Similarly, we can create a line chart to show how the probability grows over time. Run the code below!

<!-- grade_id: cell-c46ee253ad4a8777, points: 0, grade: false, solution: true -->
```python
df.plot.line(y='P(n)', x='People at The Party', grid=True)
```

### Puzzle 2.3: ❓ Re-Reflection ❓

 Knowing what you know now, how much more less likely that someone in your lab shares **your** birthday **OR** that there's at least one shared birthday between **anyone** in your lab?  How likely is it to happen in the room you're on?  (Look at the graphs and tables to get the exact probability.)

***Double click this cell and type your answer here.***

## Puzzle 3: Taylor's Birthday Problem

Taylor Swift wants **every day to be someone's birthday**!  Given a massive concert of `n` people, what is the probability that **all 365 days have at least one person with a birthday that day**?

Since Wade, Karle, and the full course staff <3s Taylor, we wrote a function to do this simulation for Taylor!  

The `birthday.forTaylor()` function returns one simulation of the number of people needed to satisfy Taylor's wish for there to be someone with a birthday for every day of the year at her concert!

You will need to **use this function** to help us calculate probabilities!

### Puzzle 3.1: Run the already-complete simulation

Run this simulation in the next cell several times.  See the different number of people at Taylor's concert for everyone to have a birthday:

<!-- grade_id: cell-a60a84fa1d64c87d, points: 0, grade: false, solution: true -->
```python
...
```

#### Collecting Many Results

Just like the previous puzzles, we have the function `birthday.forTaylor_DataFrame(n)` that returns a DataFrame for `n` **simulations**.

Collect a new DataFrame, `df_taylor`, that stores at least **1,000** simulations of Taylor's wish!

<!-- grade_id: cell-6a44b9f35daf8325, points: 0, grade: false, solution: true -->
```python
df_taylor = ...
df_taylor
```

### Puzzle 3.2: Finding the probability we delivered for Taylor

The concert is held and we have **exactly 2,000 people packed onto the quad for Taylor Swift!**  Using your simulation results, find an estimate for the probability that Taylor will be happy to know every day of the year is someone’s birthday:

<!-- grade_id: cell-0b48a78d683ce577, points: 0, grade: false, solution: true -->
```python
# First, find only the rows where Taylor's had her wish come true when there were at most 2,000 people (what conditional would that be?): 
df_success = ...
df_success
```

<!-- grade_id: cell-3cd0d620430a89ba, points: 0, grade: false, solution: true -->
```python
p_everyBirthday = ...
p_everyBirthday
```

### Puzzle 3.3: Distribution Histogram

Create a histogram of your `df_taylor` results, showing the distribution of people needed in our simulation to account for every birthday.

<!-- grade_id: cell-62147f75639a94d3, points: 0, grade: false, solution: true -->
```python
# Create a histogram:
...
```

<!-- grade_id: cell-1bae19913f8df080, points: 0, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 3 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.

assert( len(df_taylor) >= 1000 ), "Ensure you have done at least 1,000 simulations."
assert( len(df_success) > 0 ), "Ensure your success includes all simulations with 2,000 or fewer people."
assert( len(df_success) < len(df_taylor) ), "Ensure your success includes all simulations with 2,000 or fewer people."
assert( p_everyBirthday > 0.1 and p_everyBirthday < 1 ), "Ensure `p_everyBirthday` is a probability."

## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
print()
print("Great! You figured it out. Now go start inviting people!")
```

## Puzzle 4: Olivia's Birthday Problem - Everyone Shares a Birthday

Not wanting to be exactly like Taylor Swift, Olivia Rodrigo wants **everyone to share a birthday with at least one other person** AND **for every day of the year to be someone's birthday**.  We think that we will need an even bigger audience to make Olivia happy!

The `birthday.forOlivia()` function returns one simulation of the number of people needed to satisfy Olivia's wishes!

### Puzzle 4.1: Run the function

Use `birthday.forOlivia()` in the following cell and run it several times.  See the different number of people at Olivia's concert for everyone to have a birthday and share a birthday:

```python
...
```

### Puzzle 4.2: Collect Simulation Results

Identical to the previous puzzles, we have created `birthday.forOlivia_DataFrame(n)` that returns a DataFrame for `n` simulations of Olivia's wishes!

Collect the results of running the simulation at least **1000** times and store the result in `df_olivia`.

*(Note: This may take a few seconds or up to a minute to run on a slower laptop.)*

```python
df_olivia = ...
df_olivia
```

### Puzzle 4.3: Distribution Histogram

Create a histogram of your `df_olivia` results, showing the distribution of people needed in our simulation to satisfy Olivia's wishes.

```python
# Create a histogram:
...
```

### Puzzle 4.4: ❓ Estimate the Number of People Needed for Olivia's Next Concert ❓

We want the probability of Olivia being happy to know that every day of the year everyone shares their birthday with someone else to be at least **50%** for her next concert. Using the histogram from your previous puzzle, find an estimate of how many people needs to be at Olivia's next concert:<br>

**HINT:** (double-click this cell to see it): <span style="color:#ffffff00">This is the same as estimating the value associated with the 50th Percentile, or the Median.</span>

***Double click this cell and type your estimate here.***

### Puzzle 4.5: Finding the Probability For Your Estimate

Using your simulation results from Puzzle 4.2 and your estimate from Puzzle 4.3, calculate an estimate for the probability that Olivia will be happy at her next concert:

```python
# First, find the times the simulation was a success:
df_olivia_success = ...
```

```python
p_everyBirthdayShared = ...
p_everyBirthdayShared
```

```python
## == TEST CASES for Puzzle 4 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error our output, you PASSED all test cases!
# - If this cell results in any errors, check you previous cell, make changes, and RE-RUN your code and then this cell.

assert( len(df_olivia) >= 1000 ), "Ensure you have done at least 1,000 simulations."
assert( len(df_olivia_success) > 0 ), "Ensure your success includes all simulations with your median estimate or fewer people."
assert( len(df_olivia_success) < len(df_olivia) ), "Ensure your success includes all simulations with your median estimate or fewer people."
assert( p_everyBirthdayShared > 0.1 and p_everyBirthdayShared < 1 ), "Ensure `P_everyBirthdayShared` is a probability."

## == SUCCESS MESSAGE ==
# You will only see this message (with the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
print()
print("Great! You figured it out. Now go start inviting people!")
```

## Finale!

You're almost done -- congratulations!

You need to do two more things:

1. Save your work. To do this, go to **File -> Save All**
2. After you have saved, exit this notebook and follow the webpage instructions to commit this lab to your Git repository!
3. After you have updated your Git repository, check your submission in your GitHub repository online to verify it was successful!